# Scientific Programming with Python 
## Assignment 0 - 2025W

**Associate Prof. Dr. Sascha Hunold** \
Research Group for Parallel Computing \
Faculty of Informatics \
TU Wien  

---

## Task 1 (2 Points)

### Description

Write a function `summing(a_list)` that takes a list of elements as an input.
It returns the sum of all these elements from the list, which can be converted into a `float`.
For example, the string "2.3" can be converted into a float and should be considered in the sum,
but the string "bla" cannot be converted into a float and should be ignored.


### Example

For the list `[2.3, "4.5", "bla"]`, the function should return `6.8` (as a `float`).
If there is no such element, the function should return `0.0` (as a `float`).

<!-- grade_id: cell-338b48ca9733335c, points: 0, grade: false, solution: true -->
```python
def summing(a_list):
    # complete that function
    """
    Parameters
    ----------
    a_list : a list of elements
    """
    # YOUR CODE HERE
    summe = 0.0
    for i in a_list:
        try:
            converted_item = float(i)
            summe += converted_item
        except ValueError:
            summe += 0.0
    return summe
```

<!-- grade_id: cell-9c974ea537f6113d, points: 2, grade: true, solution: false -->
```python
# remember, not all tests are visible
import math
assert math.isclose(summing([-0.3, 0.3]), 0.0), summing([-0.3, 0.3])
assert math.isclose(summing([2.3, "4.5", "bla"]), 6.8), summing([2.3, "4.5", "bla"])
assert math.isclose(summing([]), 0.0), summing([])
```

## Task 2 (3 Points)

### Description

Implement a function `mins_since(a_date: str) -> int` that computes the number of minutes between a given date `a_date` and the fixed reference date **October 1, 2025, 00:00:00**.

- The argument `a_date` is a string in the format `"dd.mm.yyyy hh:mm:ss"`

where `dd` is the day, `mm` the month, `yyyy` the year, and `hh:mm:ss` the time in hours, minutes, and seconds.

- The function should return the number of **whole minutes** between `a_date` and the reference date, but **only if `a_date` is strictly earlier than the reference date**.

- If `a_date` is equal to or later than the reference date, the function should return `-1`.

- Use the calendar functionality of Python. Leap years should be considered.

### Examples

```python
mins_since("30.09.2025 23:59:00")  # → 1
mins_since("30.09.2025 23:00:00")  # → 60
mins_since("01.10.2025 12:00:00")  # → -1

<!-- grade_id: cell-e0769d98fa39802f, points: 0, grade: false, solution: true -->
```python
import calendar


def mins_since(a_date):
    # split string and get the time and date parts
    date_part, time_part = a_date.split()
    day, month, year = map(int, date_part.split('.'))
    hour, minute, second = map(int, time_part.split(':'))

    #referrence date: 01.10.2025 00:00:00
    ref_tuple = (2025, 10, 1, 0, 0, 0, 0, 0, -1)
    ref_secs = calendar.timegm(ref_tuple)  # convert to seconds (UTC)

    input_tuple = (year, month, day, hour, minute, second, 0, 0, -1)
    input_secs = calendar.timegm(input_tuple)

    # When the  date is not before 01.10.2025 00:00:00 return -1
    if input_secs >= ref_secs:
        return -1

    # Calculate the differences and divide by 60 to get the minutes
    diff_minutes = (ref_secs - input_secs) / 60
    return int(diff_minutes)

```

<!-- grade_id: cell-263cfc1a6977c77b, points: 3, grade: true, solution: false -->
```python
# remember, not all tests are visible
assert mins_since("30.09.2025 23:58:00") == 2, mins_since("30.09.2025 23:58:00")
assert mins_since("30.09.2025 23:59:00") == 1, mins_since("30.09.2025 23:59:00")
assert mins_since("29.09.2025 00:00:00") == 2880, mins_since("29.09.2025 00:00:00")
```